<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Lab-3:-Pyramid-Escape" data-toc-modified-id="Lab-3:-Pyramid-Escape-1">Lab 3: Pyramid Escape</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span></li><li><span><a href="#Code" data-toc-modified-id="Code-3">Code</a></span></li><li><span><a href="#Tests" data-toc-modified-id="Tests-4">Tests</a></span></li><li><span><a href="#Hints:" data-toc-modified-id="Hints:-5">Hints:</a></span></li></ul></div>

Lab 3: Pyramid Escape
-----

![](https://upload.wikimedia.org/wikipedia/en/thumb/0/01/Nathan_Drake_Uc4_full_body_shot.png/220px-Nathan_Drake_Uc4_full_body_shot.png)
<br>
Imagine you are Nathan Drake, Laura Croft or Indiana Jones. You have just stolen treasure from the top of a pyramid (technically, you live in [flatland](https://en.wikipedia.org/wiki/Flatland) so it is triangle). 

You have to get out the bottom of pyramid by finding the easiest path possible. Each step in the descent has an associated weight. Your goal is to find the minimum total weight of the best path.

Learning Outcomes
----

__By the end of this session, you should be able to__:

- Write a function that uses dynamic programming to find an optimal path.
- Write efficient code to solve the problem.
- Explain your logic to solve the problem.

Code
-----

```python
reset -fs
```

<!-- grade_id: cell-b40531bd4a9f86c0, points: 0, grade: false, solution: true -->
```python
"""Write the following function using dynamic programming.
"""

def min_path_weight(pyramind)-> int:
    "Find sum for minimum path through pyramind."
    
    # TODO: Write your solution below here
    

    
    
```

Tests
----

<!-- grade_id: cell-bfafcec610a29fa9, points: 1, grade: true, solution: false -->
```python
# A short pyramind. 

pyramind = [      [1],
                 [4, 3],
               [1, 5, 1],
             [1, 5, 9, 9],
            [1, 5, 9, 8, 7],
           ] 

assert min_path_weight(pyramind) == 8 
# Another way to think of this solution is: sum([1, 4, 1, 1, 1])

# Note: a greedy policy will not find the min!
assert min_path_weight(pyramind) != 21 
# Another way to think of this solution is: sum([1, 3, 1, 9, 7])
```

<!-- grade_id: cell-27fe775fa174f456, points: 1, grade: true, solution: false -->
```python
# Another short pyramind.

pyramind = [      [2],
                 [4, 3],
               [8, 5, 6],
             [4, 2, 6, 2],
            [1, 5, 2, 3, 4],
           ] 

assert min_path_weight(pyramind) == 14 == sum([2, 3, 5, 2, 2]) 
```

```python
import pickle

# The pyraminds are stored on disk as popo (plain old Python objects) in binary format to save space.
# Once loaded into memory they can printed as a lol (list of lists).
# If you print a big pyramind, you'll overload the DOM (Document Object Model)
```

<!-- grade_id: cell-7b8b982fc4709e7a, points: 1, grade: true, solution: false -->
```python
# Taller pyramind
with open("pyramind_with_10_levels", "rb") as f:
    pyramind_with_10_levels = pickle.load(f)
    
assert min_path_weight(pyramind=pyramind_with_10_levels) == 22
```

<!-- grade_id: cell-c3b49b981a692362, points: 1, grade: true, solution: false -->
```python
# Even taller pyramind
with open("pyramind_with_100_levels", "rb") as f:
    pyramind_with_100_levels = pickle.load(f)
    
assert min_path_weight(pyramind=pyramind_with_100_levels) == 193
```

<!-- grade_id: cell-2b0f93b1e3769acb, points: 1, grade: true, solution: false -->
```python
# "Now we getting somewhere with height" pyramind
with open("pyramind_with_1000_levels", "rb") as f:
    pyramind_with_1000_levels = pickle.load(f)
    
assert min_path_weight(pyramind=pyramind_with_1000_levels) == 2048 
```

<!-- grade_id: cell-68ebac116a4394d7, points: 1, grade: true, solution: false -->
```python
# Let's see if your solutions scales 😜
with open("pyramind_with_2000_levels", "rb") as f:
    pyramind_with_2000_levels = pickle.load(f)

assert min_path_weight(pyramind=pyramind_with_2000_levels) == 4057
```

```python
# Benchmark code speed
%timeit -n 1 min_path_weight(pyramind_with_2000_levels)  #=> 1.46 s ± 33.3 ms per loop (mean ± std. dev. of 7 runs, 1 loop each)
```

Output:
```
1.59 s ± 31 ms per loop (mean ± std. dev. of 7 runs, 1 loop each)
```

Hints:
------

- You can have a variable called `min_weight_to_current_row`. Print that variable out at each row to help debug.
- You might want to start with one or two level pyramid.
- On Brian's old (i.e., slow) laptop it takes less than 2 seconds to find the best path through 2,000 levels. Your code does not need to be that fast but points can be deducted for slow code. Performance is a feature!

<br>
<br> 
<br>

----